# Assignment 1
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [datasets/NISPUF17.csv](datasets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [datasets/NIS-PUF17-DUG.pdf](datasets/NIS-PUF17-DUG.pdf).

## Question 1 (20%)
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

<!-- grade_id: cell-d4625c53be146b54, points: 0, grade: false, solution: true -->
```python
import collections
def proportion_of_education():
    import pandas as pd
    df = pd.read_csv('NISPUF17.csv')
    EDUC1 = df[['EDUC1']]
    elem = EDUC1.groupby(["EDUC1"]).size()
    sum_elem = EDUC1.groupby(["EDUC1"]).size().sum()

    education_dict = {"less than high school":elem[1]/sum_elem,
                      "high school":elem[2]/sum_elem,
                    "more than high school but not college":elem[3]/sum_elem,
                    "college":elem[4]/sum_elem}
    return education_dict
proportion_of_education()
```

Output:
```
{'college': 0.47974705779026877,
 'high school': 0.172352011241876,
 'less than high school': 0.10202002459160373,
 'more than high school but not college': 0.24588090637625154}
```

<!-- grade_id: cell-75b071a5208d6122, points: 20, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
```

## Question 2 (20%)

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-2cbda7565b40aaf3, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # your code goes here
    import pandas as pd
    df = pd.read_csv('NISPUF17.csv')
    new_df = df[['P_NUMFLU','CBF_01']].dropna().groupby('CBF_01').mean()    
    tuple_influenza = (float(new_df.iloc[0]),float(new_df.iloc[1]))
    return tuple_influenza
average_influenza_doses()
```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-744f239f462ce2c9, points: 20, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."
```

## Question 3 (25%)
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated and did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value I'm looking for starts with the digits `0.00773`.

<!-- grade_id: cell-28bc178306233102, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    import pandas as pd
    import numpy as np
    df = pd.read_csv('NISPUF17.csv')
    
    df = pd.read_csv('NISPUF17.csv')
    #.loc[(df['HAD_CPOX'] == 1)]
    cpox = df[["HAD_CPOX","SEX",'P_NUMVRC']]
    #filter only by male
    male = cpox.where(cpox['SEX'] == 1).dropna() 
    #filter only by female
    female = cpox.where(cpox['SEX'] == 2).dropna() 
    
    #children with cpox
    male_had_cpox = male.where(male['HAD_CPOX'] == 1).dropna() 
    male_no_cpox = male.where(male['HAD_CPOX'] != 1).dropna()#whithout cpox
    male_yes_cpox_vaccinated = male_had_cpox.where(male_had_cpox['P_NUMVRC'] > 0).dropna()
    male_no_cpox_vaccinated = male_no_cpox.where(male_no_cpox['P_NUMVRC'] > 0).dropna()
    count_yes = male_yes_cpox_vaccinated.groupby(['HAD_CPOX','SEX']).size().to_frame("count").reset_index()
    
    count_no = male_no_cpox_vaccinated.groupby(['HAD_CPOX','SEX']).size().to_frame("count").reset_index()
    
    male_ratio = count_yes/count_no    
    male_ratio = male_ratio['count'][0]
    
    #drop not a number
    df2 = df[["HAD_CPOX","SEX",'P_NUMVRC']].dropna() 
    df3 = df2.loc[(df2["P_NUMVRC"] == 1)|(df2["P_NUMVRC"] == 2)]
    
    #df2[df2['P_NUMVRC'] != 0]  #remove the 0s
    #convert series to df
    df4 = df3.groupby(["HAD_CPOX","SEX"]).size().to_frame("count").reset_index() 

    female_cpox = (df4.loc[(df4["SEX"] == 2)&(df4["HAD_CPOX"] == 1), ["count"]]).iloc[0].item()
    female_no_cpox = (df4.loc[(df4["SEX"] == 2)&(df4["HAD_CPOX"] == 2), ["count"]]).iloc[0].item() 

    print({'male':male_ratio,'female':female_cpox / female_no_cpox    })
    dict_cpox = {'male':male_ratio,'female':female_cpox / female_no_cpox    }
    return dict_cpox
```

<!-- grade_id: cell-fba9cff0bf962f1a, points: 25, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."
```

Output:
```
{'male': 0.009675583380762664, 'female': 0.0077918259335489565}
```

## Question 4 (25%)
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. If the `had_chickenpox_column` is either `1` (for yes) or `2` for no, and that the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine, then a positive correlation (e.g. `corr > 0`) would mean that an increase in `had_chickenpox_column` (which means more no's) would mean an increase in the `num_chickenpox_vaccine_column` (which means more doses of vaccine). If `corr < 0` then there is a negative correlation, indicating that having had chickenpox is related to an increase in the number of vaccine doses. Also, `pval` refers to the probability the relationship observed is significant. In this case `pval` should be very very small (will end in `e-18` indicating a very small number), which means the result unlikely to be by chance.

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8a03b576e9dc513f, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe    
    df1 = df[["P_NUMVRC","HAD_CPOX"]].dropna()
    df2 = df1.loc[(df1["HAD_CPOX"] == 1)|(df["HAD_CPOX"] == 2)]
    df3 = df2.rename(columns= {"P_NUMVRC":"num_chickenpox_vaccine_column", "HAD_CPOX":"had_chickenpox_column"})
    
    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df3["had_chickenpox_column"],df3["num_chickenpox_vaccine_column"])
    
    # just return the correlation at the end
    # return corr
    print(corr_chickenpox())
```

```python
def corr_chickenpox():
    import scipy.stats as stats
    import pandas as pd
    import numpy as np
    df = pd.read_csv('NISPUF17.csv')
    
    df1 = df[["P_NUMVRC","HAD_CPOX"]].dropna()
    df2 = df1.loc[(df1["HAD_CPOX"] == 1)|(df["HAD_CPOX"] == 2)]
    df3 = df2.rename(columns= {"P_NUMVRC":"num_chickenpox_vaccine_column", "HAD_CPOX":"had_chickenpox_column"})
    corr, pval=stats.pearsonr(df3["had_chickenpox_column"],df3["num_chickenpox_vaccine_column"])
    return (corr)
corr_chickenpox()
```

Output:
```
0.07044873460148
```

<!-- grade_id: cell-6e19cbdf20cb9e56, points: 25, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."
```

## Question 5 (10%)

Is there a relationship between states which voted for the republican candidate for president in 2016 (Trump) versus those who voted for the democratic candidate for president in 2016 (Hillary) and the rate of vaccination in those states in 2017? I would like you to answer this by doing a correlation test between the ratio of votes for Trump/Clinton against the ratio of Vaccination/Non-Vaccination (of Chickenpox!) by state.

Note: You are certainly using joins or merges here.

Note: You will might need the state [FIPS identifiers](https://en.wikipedia.org/wiki/Federal_Information_Processing_Standard_state_code#FIPS_state_codes) and will certainly want the [MIT Elections Data](https://dataverse.harvard.edu/dataset.xhtml?persistentId=doi:10.7910/DVN/42MVDX) which has the results from 2018.

Note: I don't know the answer to this question already.

```python
#correct code
%matplotlib inline
def corr_election():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    df_president = pd.read_csv('1976-2016-president.csv')
    df_vaccine = pd.read_csv('NISPUF17.csv')


    df_president = df_president[['state','candidate','candidatevotes','state_fips','year']]

    df_president = df_president[df_president['year'] == 2016]
    df_president = (df_president.loc[df_president['candidate']
                                     .isin(['Trump, Donald J.','Clinton, Hillary'])])

    df_president = (df_president.groupby(['state_fips','candidate'])['candidatevotes']
                    .sum()
                    .unstack())

    merge_votes = df_president.reset_index()

    merge_votes = merge_votes.rename(columns= {"state_fips":"STATE"})
    merge_votes['ratio'] = merge_votes['Trump, Donald J.'] / merge_votes['Clinton, Hillary'] 

    df_vaccine = pd.read_csv('NISPUF17.csv', delimiter = ',')
    state_vaccine = df_vaccine[["STATE","P_NUMVRC"]].dropna()
    no_vac = (state_vaccine[state_vaccine.P_NUMVRC == 0]
              .groupby(["STATE"])
              .count())

    yes_vac = (state_vaccine[(state_vaccine.P_NUMVRC >= 1) 
                             & (state_vaccine.P_NUMVRC < 77 )]
               .groupby(["STATE"])
               .count())

    merge_vac = (yes_vac
                 .merge(no_vac, left_index=True, right_index=True)
                 .reset_index())
    merge_vac['ratio'] = merge_vac['P_NUMVRC_x'] / merge_vac['P_NUMVRC_y']


    corr, pval=stats.pearsonr(merge_votes['ratio'],merge_vac['ratio'])

    corr
    return corr, pval
corr_election()
```

Output:
```
(-0.28010835423617825, 0.04649924883873912)
```

```python
#alternate answer

%matplotlib inline
import scipy.stats as stats
import numpy as np
import pandas as pd
df_president = pd.read_csv('1976-2016-president.csv')
df_vaccine = pd.read_csv('NISPUF17.csv')

#voting data
df_president = (df_president[['state_fips','candidate','candidatevotes','party','writein']]
                .dropna()
                .loc[df_president['year'] == 2016]) #get useful columns and year 2016
trump_votes = (df_president
               .loc[(df_president['candidate'] == 'Trump, Donald J.')& ~(df_president["writein"])]
               .loc[df_president['party'] == 'republican']) #Get trump, clean up party column and writein column
clinton_votes = (df_president
                 .loc[(df_president['candidate'] == 'Clinton, Hillary')& ~(df_president["writein"])]
                 .loc[df_president['party'] == 'democrat']) #likewise
trump_clinton_votes = pd.merge(trump_votes, clinton_votes, on = "state_fips") #merge
trump_clinton_votes['ratio'] = (trump_clinton_votes['candidatevotes_x'])/trump_clinton_votes['candidatevotes_y'] #ratio

#vaccine data
df_vaccine = df_vaccine[["STATE","P_NUMVRC"]]#get two useful columns
df_no_vaccine = (df_vaccine [df_vaccine['P_NUMVRC'] == 0]
                 .groupby(["STATE"])
                 .size()
                 .to_frame("count")
                 .reset_index()
                 .rename(columns= {"STATE":"state", "count":"no_vaccine"})) #get no vaccine data
df_yes_vaccine = (df_vaccine [df_vaccine['P_NUMVRC'] >= 1]
                  .groupby(["STATE"])
                  .size()
                  .to_frame("count")
                  .reset_index()
                  .rename(columns= {"STATE":"state", "count":"yes_vaccine"})) #get yes vaccine data
merge_vaccine = pd.merge(df_no_vaccine,df_yes_vaccine , on = "state") #merge
ratio_vaccine = merge_vaccine['yes_vaccine']/(( merge_vaccine['no_vaccine']))
corr, pval=stats.pearsonr(trump_clinton_votes['ratio'],ratio_vaccine)
corr
```

Output:
```
-0.2790809695271168
```

```python

```